---
title: 'Bases de Python'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

explicaci칩n de la funci칩n print

`@instructions`
tienes que imprimir las palabras "Hola Mundo" utilizando la funci칩n print(), simplemente colocando dentro del par치ntesis las palabras "Hola Mundo"

`@hint`


`@pre_exercise_code`
```{python}
# vamos a comenzar con el primer programa 
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("hola mundo")
```

`@sct`
```{python}

```
