---
title: 'Bases de Python'
description: 'Chapter description goes here.'
free_preview: true
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

explicación de la función print

`@instructions`
tienes que imprimir la palabra

`@hint`


`@pre_exercise_code`
```{python}
# vamos a comenzar con el primer programa 
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("hola mundo")
```

`@sct`
```{python}

```
