---
title: 'Bases de Python'
description: 'Chapter description goes here.'
free_preview: true
---

## Hola mundo

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

explicaci칩n de la funci칩n print

`@instructions`
tienes que imprimir las palabras "Hola Mundo" utilizando la funci칩n print(), simplemente colocando dentro del par치ntesis las palabras "Hola Mundo"

`@hint`


`@pre_exercise_code`
```{python}
# vamos a comenzar con el primer programa 
```

`@sample_code`
```{python}

```

`@solution`
```{python}
print("hola mundo")
```

`@sct`
```{python}

```

---

## Insert exercise title here

```yaml
type: VideoExercise
key: dc95f3fd6e
xp: 50
```

`@projector_key`
44abc4f6df15242e9b2e2b6e8bbecb0f

---

## Insert exercise title here

```yaml
type: TabExercise
key: 92f571673e
xp: 100
```



`@pre_exercise_code`
```{python}

```
